# Exercise 2 - shell, pipes, and csvkit

In this exercise, we'll review a few shell commands and explore working with pipes and csvkit.

You will need to fill in commands whenever prompted.  Please replace the text with your solution.

Remember to submit your completed `.ipynb` file to Blackboard and to add/commit it to your Git repository and push it to GitHub.


## Part 1 - shell commands, redirection, and pipes

### Basic shell commands and redirection

Create a directory called `part1` using `mkdir`.

```python
!mkdir part1
```

Rename `part1` to `partone` using `mv`.

```python
!mv part1 partone
!ls
```

Output:
```
exercise-01.ipynb exercise-02.ipynb [34mpartone[m[m
```

Create a file named `filelist.txt` using the output from `ls` and the output redirector `>`.

```python
!ls > filelist.txt
```

```python
!cat filelist.txt
```

Output:
```
exercise-01.ipynb
exercise-02.ipynb
filelist.txt
partone
```

Append to `filelist.txt` using the output appending redirector `>>`.  Note the difference between the single `>` and double `>>`.

```python
!ls >> filelist.txt
!cat filelist.txt
```

Output:
```
exercise-01.ipynb
exercise-02.ipynb
filelist.txt
partone
exercise-01.ipynb
exercise-02.ipynb
filelist.txt
partone
```

```python
!ls > filelist.txt
!cat filelist.txt
```

Output:
```
exercise-01.ipynb
exercise-02.ipynb
filelist.txt
partone
```

What's the difference between `>` and `>>`?


<!-- grade_id: 1060, points: 1, grade: true, solution: true -->
Both >> and > redirect outputs of a command to a file. If the file does not exists in the path, they create the file. Otherwise,  > overwrites the file. On the other hand, >> appends the outputs to the end of the file.

### Your turn

Complete the following tasks in the cells provided.  All the tests in the testing cells (with the `assert` statements) below should pass without error - be sure to execute those as well, and if you see errors, fix your answer and try testing again until there are no errors.

Create a directory called `mydirectory`.

<!-- grade_id: 1070, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
!mkdir mydirectory 
```

<!-- grade_id: 2080, points: 1, grade: true, solution: false -->
```python
import os
assert 'mydirectory' in os.listdir('.')
```

Using `ls` and output redirection, create a file called `myfiles.txt` in the directory `mydirectory` that contains the list of files in the current directory.

<!-- grade_id: 1080, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
!ls > ./mydirectory/myfiles.txt
```

<!-- grade_id: 2082, points: 1, grade: true, solution: false -->
```python
assert 'myfiles.txt' in os.listdir('mydirectory')
```

<!-- grade_id: 2084, points: 1, grade: true, solution: false -->
```python
myfiles = open('mydirectory/myfiles.txt').read()
assert 'exercise-02.ipynb' in myfiles
```

Clean up the directory you just created by removing its contents (the file you created) using `rm`.

<!-- grade_id: 2086, points: 1, grade: true, solution: false -->
```python
!rm ./mydirectory/myfiles.txt
```

<!-- grade_id: 2088, points: 1, grade: true, solution: false -->
```python
assert 'myfiles.txt' not in os.listdir('mydirectory')
```

Now remove the directory itself using `rmdir`.

<!-- grade_id: 2090, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
!rmdir mydirectory/
```

<!-- grade_id: 2092, points: 1, grade: true, solution: false -->
```python
assert 'mydirectory' not in os.listdir('.')
```

### Filters and pipes

Let's look at something a little more interesting.  Download the text of Herman Hesse's *Siddhartha* from [Project Gutenberg](http://www.gutenberg.org/):

```python
!wget https://github.com/gwsb-istm-6212-fall-2016/syllabus-and-schedule/raw/master/exercises/pg2500.txt
```

Output:
```
--2016-09-07 09:51:25--  https://github.com/gwsb-istm-6212-fall-2016/syllabus-and-schedule/raw/master/exercises/pg2500.txt
Resolving github.com... 192.30.253.112
Connecting to github.com|192.30.253.112|:443... connected.
HTTP request sent, awaiting response... 302 Found
Location: https://raw.githubusercontent.com/gwsb-istm-6212-fall-2016/syllabus-and-schedule/master/exercises/pg2500.txt [following]
--2016-09-07 09:51:26--  https://raw.githubusercontent.com/gwsb-istm-6212-fall-2016/syllabus-and-schedule/master/exercises/pg2500.txt
Resolving raw.githubusercontent.com... 151.101.32.133
Connecting to raw.githubusercontent.com|151.101.32.133|:443... connected.
HTTP request sent, awaiting response... 200 OK
Length: 241176 (236K) [text/plain]
Saving to: ‘pg2500.txt’

pg2500.txt          100%[===================>] 235.52K  1.42MB/s    in 0.2s    

2016-09-07 09:51:27 (1.42 MB/s) - ‘pg2500.txt’ saved [241176/241176]
```

*Note*: sometimes Project Gutenberg restricts access.  If that creates an error for you, you should be able to `wget` the same file from our class repository on GitHub at the url https://github.com/gwsb-istm-6212-fall-2016/syllabus-and-schedule/raw/master/exercises/pg2500.txt.

However you get the file, let's rename it to something easier to remember.

```python
!mv pg2500.txt siddhartha.txt
```

`head` and `tail` are very useful.  They let you take a quick peek at the start and end of files.

```python
!head siddhartha.txt
```

Output:
```
﻿The Project Gutenberg EBook of Siddhartha, by Herman Hesse

This eBook is for the use of anyone anywhere at no cost and with
almost no restrictions whatsoever.  You may copy it, give it away or
re-use it under the terms of the Project Gutenberg License included
with this eBook or online at www.gutenberg.org


Title: Siddhartha
```

```python
!tail siddhartha.txt
```

Output:
```
Most people start at our Web site which has the main PG search facility:

     http://www.gutenberg.org

This Web site includes information about Project Gutenberg-tm,
including how to make donations to the Project Gutenberg Literary
Archive Foundation, how to help produce our new eBooks, and how to
subscribe to our email newsletter to hear about new eBooks.
```

`grep` is one of the most useful filters.  It lets you search for and match lines that contain specific expressions.  For example, to find mentions of "copyright":

```python
!grep copyright siddhartha.txt
```

Output:
```
one owns a United States copyright in these works, so the Foundation
permission and without paying copyright royalties.  Special rules,
(trademark/copyright) agreement.  If you do not agree to abide by all
or PGLAF), owns a compilation copyright in the collection of Project
1.D.  The copyright laws of the place where you are located also govern
the copyright status of any work in any country outside the United
posted with permission of the copyright holder), the work can be copied
with the permission of the copyright holder, your use and distribution
terms imposed by the copyright holder.  Additional terms will be linked
permission of the copyright holder found at the beginning of this work.
effort to identify, do copyright research on, transcribe and proofread
corrupt data, transcription errors, a copyright or other intellectual
unless a copyright notice is included.  Thus, we do not necessarily
```

Notice anything that those lines have in common?

Let's add a little more information by including the `-n` flag to add matching line numbers.

```python
!grep -n copyright siddhartha.txt
```

Output:
```
3973:one owns a United States copyright in these works, so the Foundation
3975:permission and without paying copyright royalties.  Special rules,
4010:(trademark/copyright) agreement.  If you do not agree to abide by all
4029:or PGLAF), owns a compilation copyright in the collection of Project
4044:1.D.  The copyright laws of the place where you are located also govern
4051:the copyright status of any work in any country outside the United
4070:posted with permission of the copyright holder), the work can be copied
4080:with the permission of the copyright holder, your use and distribution
4082:terms imposed by the copyright holder.  Additional terms will be linked
4084:permission of the copyright holder found at the beginning of this work.
4155:effort to identify, do copyright research on, transcribe and proofread
4160:corrupt data, transcription errors, a copyright or other intellectual
4308:unless a copyright notice is included.  Thus, we do not necessarily
```

Now let's look for any mention of "river".  This will match a lot of text, so we'll just take the first 10 matching lines by *piping* the output from `grep` into `head`.

```python
!grep -n river siddhartha.txt | head
```

Output:
```
52:In the shade of the house, in the sunshine of the riverbank near the
56:tanned his light shoulders by the banks of the river when bathing,
107:came into his mind, flowing from the water of the river, sparkling from
329:the river and to perform the first ablution."
460:the same short numbing is what the driver of an ox-cart finds in the
469:is no driver of an ox-cart and a Samana is no drunkard.  It's true that
580:submerged in the murky river of physical forms.  Many wonderful and
1062:and the river flowed, the forest and the mountains were rigid, all of it
1065:all this yellow and blue, river and forest, entered Siddhartha for the
1069:who scorns diversity, who seeks unity.  Blue was blue, river was river,
```

How many lines contain "river"?  We can count by piping into the word count tool `wc`.

```python
!grep river siddhartha.txt | wc
```

Output:
```
109    1365    7443
```

That's 109 matching lines, containing 1365 words and 7979 characters.  If you just wanted the lines by themselves, use `wc -l`:

```python
!grep river siddhartha.txt | wc -l
```

Output:
```
109
```

What if we want to match both upper- and lower-case text?  Use `grep -i`:

```python
!grep time siddhartha.txt | wc -l
```

Output:
```
166
```

```python
!grep -i time siddhartha.txt | wc -l
```

Output:
```
167
```

### Your turn

How many lines in *Siddhartha* contain "other" (just lower-case)?  Start by using `grep` to extract lines that match the word "other" in `siddhartha.txt` and redirecting it to a file called `other-lines.txt`.

<!-- grade_id: 2180, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
!grep other siddhartha.txt > other-lines.txt
```

<!-- grade_id: 2190, points: 1, grade: true, solution: false -->
```python
%sc h_other = head -1 other-lines.txt
assert "other" in h_other
```

<!-- grade_id: 2200, points: 1, grade: true, solution: false -->
```python
%sc t_other = tail -1 other-lines.txt
assert "other" in t_other
```

Now count up the lines in the file you created using `wc`.

<!-- grade_id: 2210, points: 0, grade: false, solution: true -->
```python
# it seems _ does not hold outputs of last executed statement in the notebook. 
# I use the shell capture magic here.
%sc  _ = wc -l other-lines.txt
```

<!-- grade_id: 2220, points: 1, grade: true, solution: false -->
```python
assert "127" in _
```

Your answer should be 127!

### Counting words with `grep`

By piping commands together we can do a lot of powerful things right at the command line.  Let's create a count of the most commonly occurring words in *Siddhartha*.  To do that, we could write a Python script that just counts words, but with the command line shell tools we only need to put a proper pipeline together and we can often accomplish tasks like this in one line.

First we need to split up the text lines into a word per line.  There are `grep` flags for that!

```python
!cat siddhartha.txt | grep -oE '\w{{2,}}' | head -10
```

Output:
```
The
Project
Gutenberg
EBook
of
Siddhartha
by
Herman
Hesse
This
```

Now we need to sort them and count the unique tokens.  `sort` solves the first problem.

```python
!cat siddhartha.txt | grep -oE '\w{{2,}}' | sort | head -10
```

Output:
```
000
1500
1887
20
2001
2008
2011
2013
23
2500
sort: write failed: standard output: Broken pipe
sort: write error
```

And `uniq -c` solves the second problem.

```python
!cat siddhartha.txt | grep -oE '\w{{2,}}' | sort | uniq -c | head -25
```

Output:
```
1 000
   1 1500
   1 1887
   1 20
   1 2001
   1 2008
   1 2011
   1 2013
   1 23
   4 2500
   1 30
   1 4557
   1 50
   2 501
   1 596
   1 60
   1 6221541
   1 64
   1 801
   1 809
   1 84116
   2 90
   1 99712
   1 ACTUAL
   2 AGREE
```

But there's a catch... do you see it?

We need to convert all the words down into lower case so that we are correctly counting unique words.  There's another command, `tr`, for that.

```python
!cat siddhartha.txt | grep -oE '\w{{2,}}' | tr '[:upper:]' '[:lower:]' | sort | uniq -c | head -25
```

Output:
```
1 000
   1 1500
   1 1887
   1 20
   1 2001
   1 2008
   1 2011
   1 2013
   1 23
   4 2500
   1 30
   1 4557
   1 50
   2 501
   1 596
   1 60
   1 6221541
   1 64
   1 801
   1 809
   1 84116
   2 90
   1 99712
   2 abandon
   2 abandoned
```

...and if we want to know only the top 10 words in Siddhartha, we need to sort the output.

```python
!cat siddhartha.txt | grep -oE '\w{{2,}}' | tr '[:upper:]' '[:lower:]' | sort | uniq -c | sort | head -10
```

Output:
```
1 000
   1 1500
   1 1887
   1 20
   1 2001
   1 2008
   1 2011
   1 2013
   1 23
   1 30
```

But that sorts by character, not number.  Fortunately, `sort -n` does what we want.

```python
!cat siddhartha.txt | grep -oE '\w{{2,}}' | tr '[:upper:]' '[:lower:]' | sort | uniq -c | sort -n | head -10
```

Output:
```
1 000
   1 1500
   1 1887
   1 20
   1 2001
   1 2008
   1 2011
   1 2013
   1 23
   1 30
```

But that's the wrong end of the list!  Two ways to fix that:  (a) use `tail` instead of `head`; (b) use `sort -rn`, which will sort in reverse order.  Let's try the latter.

```python
!cat siddhartha.txt | grep -oE '\w{{2,}}' | tr '[:upper:]' '[:lower:]' | sort | uniq -c | sort -rn | head -10
```

Output:
```
2221 the
1434 and
1225 to
1106 of
 960 he
 708 his
 686 in
 540 you
 524 had
 512 was
```

### Your turn

Download *Alice in Wonderland* from http://www.gutenberg.org/cache/epub/11/pg11.txt (or https://github.com/gwsb-istm-6212-fall-2016/syllabus-and-schedule/raw/master/exercises/pg11.txt if the first url doesn't work).

<!-- grade_id: 2290, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
!wget https://github.com/gwsb-istm-6212-fall-2016/syllabus-and-schedule/raw/master/exercises/pg11.txt 
```

Output:
```
--2016-09-07 10:40:18--  https://github.com/gwsb-istm-6212-fall-2016/syllabus-and-schedule/raw/master/exercises/pg11.txt
Resolving github.com... 192.30.253.112
Connecting to github.com|192.30.253.112|:443... connected.
HTTP request sent, awaiting response... 302 Found
Location: https://raw.githubusercontent.com/gwsb-istm-6212-fall-2016/syllabus-and-schedule/master/exercises/pg11.txt [following]
--2016-09-07 10:40:18--  https://raw.githubusercontent.com/gwsb-istm-6212-fall-2016/syllabus-and-schedule/master/exercises/pg11.txt
Resolving raw.githubusercontent.com... 151.101.32.133
Connecting to raw.githubusercontent.com|151.101.32.133|:443... connected.
HTTP request sent, awaiting response... 200 OK
Length: 167518 (164K) [text/plain]
Saving to: ‘pg11.txt’

pg11.txt            100%[===================>] 163.59K  --.-KB/s    in 0.05s   

2016-09-07 10:40:20 (3.21 MB/s) - ‘pg11.txt’ saved [167518/167518]
```

<!-- grade_id: 2295, points: 1, grade: true, solution: false -->
```python
assert 'pg11.txt' in os.listdir('.')
```

Now rename `pg11.txt` to `alice.txt`.

<!-- grade_id: 2300, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
!mv pg11.txt alice.txt
```

<!-- grade_id: 2310, points: 1, grade: true, solution: false -->
```python
assert 'alice.txt' in os.listdir('.')
```

```python
!cat siddhartha.txt | grep -oE '\w{{2,}}' 

```

Output:
```
The
Project
Gutenberg
EBook
of
Siddhartha
by
Herman
Hesse
This
eBook
is
for
the
use
of
anyone
anywhere
at
no
cost
and
with
almost
no
restrictions
whatsoever
You
may
copy
it
give
it
away
or
re
use
it
under
the
terms
of
the
Project
Gutenberg
License
included
with
this
eBook
or
online
at
www
gutenberg
org
Title
Siddhartha
Author
Herman
Hesse
Translator
Gunther
Olesch
Anke
Dreher
Amy
Coulter
Stefan
Langer
and
Semyon
Chaichenets
Release
Date
April
2008
EBook
2500
Last
updated
July
2011
Last
updated
January
23
2013
Language
English
START
OF
THIS
PROJECT
GUTENBERG
EBOOK
SIDDHARTHA
Produced
by
Michael
Pullen
Chandra
Yenco
Isaac
Jones
SIDDHARTHA
An
Indian
Tale
by
Hermann
Hesse
FIRST
PART
To
Romain
Rolland
my
dear
friend
THE
SON
OF
THE
BRAHMAN
In
the
shade
of
the
house
in
the
sunshine
of
the
riverbank
near
the
boats
in
the
shade
of
the
Sal
wood
forest
in
the
shade
of
the
fig
tree
is
where
Siddhartha
grew
up
the
handsome
son
of
the
Brahman
the
young
falcon
together
with
his
friend
Govinda
son
of

```

Take a look at the next cell.  Will it find the top 25 unique words in *Alice in Wonderland* successfully?

```python
!cat alice.txt | grep -oE '\w{{2,}}' | sort | uniq -c | head -25
```

Output:
```
1 000
   4 11
   1 1500
   1 1887
   1 1994
   2 20
   1 2001
   1 2008
   1 2011
   1 25
   1 30
   1 4557
   1 50
   2 501
   1 596
   1 60
   1 6221541
   1 64
   1 801
   1 809
   1 84116
   2 90
   1 99712
   1 ACTUAL
   3 ADVENTURES
```

Describe what needs to be done to the previous cell to get it to work correctly.  Describe it using words, explaining the issues, rather than using shell commands!

<!-- grade_id: 2330, points: 1, grade: true, solution: true -->
We have to sort reversely the outputs of "uniq -c" by the count of words before using "head -25" to retrieve the top 25 unique words. And We have to transform uppercase letters to lowercase after the grep command. We want to put the transform command before the first sort command because we want to reduce number of words that the sort command has to process. 

Okay, now implement your solution using shell commands with a pipeline.

<!-- grade_id: 2340, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
!cat alice.txt | grep -oE '\w{{2,}}' | tr '[:upper:]' '[:lower:]' | sort | uniq -c | sort -rn | head -25
```

Output:
```
1818 the
 940 and
 809 to
 631 of
 610 it
 553 she
 481 you
 462 said
 431 in
 403 alice
 358 was
 330 that
 274 as
 248 her
 228 with
 227 at
 204 on
 200 all
 181 this
 179 for
 178 had
 175 but
 167 be
 166 not
 155 they
```

## Part 2 - csvkit basics

Let's look at some CSV data using csvkit.  Download the 2015 fourth quarter trip dataset from [Capital Bikeshare's trip history data](https://www.capitalbikeshare.com/trip-history-data):

```python
!wget https://s3.amazonaws.com/capitalbikeshare-data/2015-Q4-cabi-trip-history-data.zip
```

Output:
```
--2016-09-07 11:11:20--  https://s3.amazonaws.com/capitalbikeshare-data/2015-Q4-cabi-trip-history-data.zip
Resolving s3.amazonaws.com... 52.216.225.107
Connecting to s3.amazonaws.com|52.216.225.107|:443... connected.
HTTP request sent, awaiting response... 200 OK
Length: 14308794 (14M) [application/zip]
Saving to: ‘2015-Q4-cabi-trip-history-data.zip’

2015-Q4-cabi-trip-h 100%[===================>]  13.65M  3.76MB/s    in 3.8s    

2016-09-07 11:11:25 (3.63 MB/s) - ‘2015-Q4-cabi-trip-history-data.zip’ saved [14308794/14308794]
```

Let's unzip it, rename it to something short, and take a look.

```python
!unzip 2015-Q4-cabi-trip-history-data.zip
```

Output:
```
Archive:  2015-Q4-cabi-trip-history-data.zip
  inflating: 2015-Q4-Trips-History-Data.csv
```

```python
!mv 2015-Q4-Trips-History-Data.csv q4.csv
```

```python
!head q4.csv
```

Output:
```
Duration (ms),Start date,End date,Start station number,Start station,End station number,End station,Bike #,Member type
166050,10/1/2015 0:01,10/1/2015 0:04,31602,Park Rd & Holmead Pl NW,31105,14th & Harvard St NW,W21109,Registered
379172,10/1/2015 0:01,10/1/2015 0:07,31314,34th & Water St NW,31237,25th St & Pennsylvania Ave NW,W20603,Registered
696038,10/1/2015 0:01,10/1/2015 0:13,31214,17th & Corcoran St NW,31214,17th & Corcoran St NW,W01233,Registered
219423,10/1/2015 0:02,10/1/2015 0:06,31104,Adams Mill & Columbia Rd NW,31121,Calvert St & Woodley Pl NW,W00218,Registered
253230,10/1/2015 0:03,10/1/2015 0:07,31102,11th & Kenyon St NW,31102,11th & Kenyon St NW,W21612,Registered
655251,10/1/2015 0:03,10/1/2015 0:14,31242,18th St & Pennsylvania Ave NW,31114,18th St & Wyoming Ave NW,W22093,Registered
309212,10/1/2015 0:06,10/1/2015 0:11,31280,11th & S St NW,31278,18th & R St NW,W00231,Registered
776195,10/1/2015 0:07,10/1/2015 0:20,31226,34th St & Wisconsin Ave NW,31214,17
```

csvkit gives us great tools for examining and working with CSV data.  We start by looking at the columns:

```python
!csvcut -n q4.csv
```

Output:
```
1: Duration (ms)
  2: Start date
  3: End date
  4: Start station number
  5: Start station
  6: End station number
  7: End station
  8: Bike #
  9: Member type
```

We can also extract just a few columns with `csvcut`:

```python
!csvcut -c1,5,7 q4.csv | head -10
```

Output:
```
Duration (ms),Start station,End station
166050,Park Rd & Holmead Pl NW,14th & Harvard St NW
379172,34th & Water St NW,25th St & Pennsylvania Ave NW
696038,17th & Corcoran St NW,17th & Corcoran St NW
219423,Adams Mill & Columbia Rd NW,Calvert St & Woodley Pl NW
253230,11th & Kenyon St NW,11th & Kenyon St NW
655251,18th St & Pennsylvania Ave NW,18th St & Wyoming Ave NW
309212,11th & S St NW,18th & R St NW
776195,34th St & Wisconsin Ave NW,17th & Corcoran St NW
151604,Wilson Blvd & N Uhle St,N Veitch  & 20th St N
```

...and make it look better with `csvlook`:

```python
!csvcut -c1,5,7 q4.csv | head -10 | csvlook
```

Output:
```
|----------------+-------------------------------+--------------------------------|
|  Duration (ms) | Start station                 | End station                    |
|----------------+-------------------------------+--------------------------------|
|  166050        | Park Rd & Holmead Pl NW       | 14th & Harvard St NW           |
|  379172        | 34th & Water St NW            | 25th St & Pennsylvania Ave NW  |
|  696038        | 17th & Corcoran St NW         | 17th & Corcoran St NW          |
|  219423        | Adams Mill & Columbia Rd NW   | Calvert St & Woodley Pl NW     |
|  253230        | 11th & Kenyon St NW           | 11th & Kenyon St NW            |
|  655251        | 18th St & Pennsylvania Ave NW | 18th St & Wyoming Ave NW       |
|  309212        | 11th & S St NW                | 18th & R St NW                 |
|  776195        | 34th St & Wisconsin Ave NW    | 17th & Corcoran St NW          |
|  151604        | Wilson Blvd & N Uhle St       | N Veitch  & 20
```

It gets even better.  Try `csvgrep`:

```python
!csvcut -c1,5,7 q4.csv | csvgrep -c3 -m '21st & I St NW' | head -10 | csvlook
```

Output:
```
|----------------+----------------------------------+-----------------|
|  Duration (ms) | Start station                    | End station     |
|----------------+----------------------------------+-----------------|
|  872420        | 5th St & Massachusetts Ave NW    | 21st & I St NW  |
|  1249207       | 11th & Kenyon St NW              | 21st & I St NW  |
|  548066        | 17th & Corcoran St NW            | 21st & I St NW  |
|  604362        | 18th St & Wyoming Ave NW         | 21st & I St NW  |
|  1556978       | Army Navy Dr & S Nash St         | 21st & I St NW  |
|  586738        | New Hampshire Ave & T St NW      | 21st & I St NW  |
|  259421        | 17th & K St NW / Farragut Square | 21st & I St NW  |
|  269179        | M St & Pennsylvania Ave NW       | 21st & I St NW  |
|  227962        | 25th St & Pennsylvania Ave NW    | 21st & I St NW  |
|----------------+----------------------------------+-----------------|
```

But wait, there's more:

```python
!csvcut -c1,5,7 q4.csv | csvgrep -c3 -m '21st & I St NW' | csvsort -c2 | head -10 | csvlook
```

Output:
```
|----------------+-----------------------+-----------------|
|  Duration (ms) | Start station         | End station     |
|----------------+-----------------------+-----------------|
|  890758        | 10th & E St NW        | 21st & I St NW  |
|  1074547       | 10th & E St NW        | 21st & I St NW  |
|  751900        | 10th & E St NW        | 21st & I St NW  |
|  938747        | 10th & Florida Ave NW | 21st & I St NW  |
|  1001635       | 10th & Florida Ave NW | 21st & I St NW  |
|  837341        | 10th & Florida Ave NW | 21st & I St NW  |
|  891224        | 10th & Florida Ave NW | 21st & I St NW  |
|  1070924       | 10th & Florida Ave NW | 21st & I St NW  |
|  799070        | 10th & Florida Ave NW | 21st & I St NW  |
|----------------+-----------------------+-----------------|
```

And you can perform basic statistics very easily:

```python
!csvcut -c1,5,7 q4.csv | csvgrep -c3 -m '21st & I St NW' | csvcut -c1 | csvstat
```

Output:
```
1. Duration (ms)
	<class 'int'>
	Nulls: False
	Min: 65972
	Max: 29254066
	Sum: 3721091088
	Mean: 867387.2
	Median: 678586.5
	Standard Deviation: 1148902.3949097062
	Unique values: 4281
	5 most frequent values:
		889715:	2
		948623:	2
		716896:	2
		134297:	2
		586081:	2

Row count: 4290
```

### Your turn

Which set of trips had the longer average trip duration:  trips *starting* at "Massachusetts Ave & Dupont Circle NW", or trips *ending* at "Massachusetts Ave & Dupont Circle NW"?

Use as many new cells as you need to compute the answer, and then write in your answer below in the "YOUR ANSWER HERE" cell.

<!-- grade_id: 2450, points: 3, grade: true, solution: true -->
```python
# the average duration of trips starting at "Massachusetts Ave & Dupont Circle NW"
!csvcut -c1,5 q4.csv | csvgrep -c2 -m "Massachusetts Ave & Dupont Circle NW" | csvcut -c1 | csvstat | grep Mean
```

Output:
```
Mean: 946431.0750804995
```

```python
# the average duration of trips ending at "Massachusetts Ave & Dupont Circle NW"
!csvcut -c1,7 q4.csv | csvgrep -c2 -m "Massachusetts Ave & Dupont Circle NW" | csvcut -c1 | csvstat | grep Mean
```

Output:
```
Mean: 840543.3018618869
```

Trips starting at "Massachusetts Ave & Dupont Circle NW" had the longer average trip duration than  trips ending at "Massachusetts Ave & Dupont Circle NW". 